# Assignment 1


**Deadline:** Friday 13th November, 10a.m. London Time

Be sure to read and understand the file ['Instructions'](./Instructions.md) for detailed information on this assignment.


**IMPORTANT: Please type the following sentence in the next cell: "I agree to solve this assignment according to the given instructions and UCL's ethics code."**




<!-- grade_id: 8c8876, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

```python
# Some packages you might need in the following

from nose.tools import assert_equal, assert_raises  # Some functions to produce some of the checks before

import numpy as np
import matplotlib.pylab as plt
from numpy.random import default_rng
from scipy.optimize import minimize


```

## Question 1. [40 points]

An investor has utility function $$u(x) = \begin{cases} \frac{|x|^{p}}{x} & \text{ if } x\neq 0;\\ 0 & \text{ if } x=0\end{cases},$$ for $p>1$.

- 1.1. Implement this utility function [10 points]

<!-- grade_id: q11, points: 0, grade: false, solution: true -->
```python
def u(x, p): 
    """
    Compute the utility function $u$ defined above on a real 'x' and a given fixed value 'p'>1
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: v11values, points: 6, grade: true, solution: false -->
```python
# Check that the definition is correct for some inputs:

assert_equal(u(0,2), 0)
assert_equal(u(2,2), 2)

```

<!-- grade_id: v11errors, points: 4, grade: true, solution: false -->
```python

```

- 1.2. Define now a function 'u_vec' that receives a numpy array or a list and a real p larger than one and returns an array with the utility function above applied entry-wise [10 points].

<!-- grade_id: q12, points: 0, grade: false, solution: true -->
```python
def u_vec(xvec, p):  
    """
    Receives a numpy array or a list 'xvec' and a real 'p'>1 and returns an array with the utility function above applied entry-wise
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: v12, points: 10, grade: true, solution: false -->
```python

np.testing.assert_array_equal(u_vec([-1,0,2],2),np.array([-1,0,2]))
assert_raises(TypeError, u_vec, ['a'],2  )

```

- 1.3 Create a function that calculates the coefficient of absolute risk aversion of this function for a given real 'x' [10 points]


<!-- grade_id: q13, points: 0, grade: false, solution: true -->
```python
def ara(x,p):
    """
    Compute the coefficient of absolute risk aversion 
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: v13a, points: 7, grade: true, solution: false -->
```python
# Check that the definition is correct for some inputs:

assert_equal(ara(1,2), 0)
assert_equal(ara(2,3), -1/2)

```

<!-- grade_id: v13b, points: 3, grade: true, solution: false -->
```python

assert_raises(ValueError,ara, 0,3)

```

- 1.4. Comment on the risk behaviour of an investor with this function [10 points]


<!-- grade_id: v14, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 2. [20 points]

Consider a market model in one-period with three assets: a risk free asset with unit return, and two risky assets with log normal returns

$R^1_1\sim \exp(N_1), \qquad R^2_1\sim \exp(N_2)$
where $N_1\sim \mathcal N(m_1,\sigma_1^2)$, $N_2\sim \mathcal N(m_2,\sigma_2^2) $ and the Gaussian random variables are correlated $$\mathrm{cor}(N_1,N_2)=\rho.$$ 



- 2.1. Explain if this market is complete [10 points]


<!-- grade_id: v21, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE

- 2.2. Create a function that returns a sample of a given size of (${R}^1_1$, ${R}_1^2$). [10 points]


<!-- grade_id: q22, points: 0, grade: false, solution: true -->
```python
def make_sample(n, m1,m2, sigma1,sigma2,rho): 
    """
    Create a sample of (${R}^1_1$, ${R}_1^2$) of size 'n', where the other parameters are defined as given. Returns a tuple (R1_sample,R2_sample)
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    
    
    
```

<!-- grade_id: v22, points: 10, grade: true, solution: false -->
```python

assert_equal(type(make_sample(100,1e-2,1e-2,1e-2,1e-2,0)),tuple)
assert_equal(len(make_sample(100,1e-2,1e-2,1e-2,1e-2,0)[0]),len(make_sample(100,1e-2,1e-2,1e-2,1e-2,0)[1]))

```

## Question 3 [40 points]

- 3.1. Let $\pi$ be a strategy (in terms of percentages) in a market with three assets (one risk-free and two risky as in Question 2). Let $\hat \pi$ be the components of the strategy on the risky assets. Define a function that returns the expected utility for an investor with a utility function as in Question 1 and initial wealth $w_0$, after applying the strategy represented by $\hat \pi$ (recall that the investment on the risk-less asset can be deduced from the budget constraint). The function should also receive a sample of $\hat R_1$. [20 points]

<!-- grade_id: q31, points: 0, grade: false, solution: true -->
```python
def u_exp(pi_hat, w0,R_hat,p):  
    """
    Returns the expected utility of Q1 (with parameter 'p') of the strategy characterised by 'pi_hat', from a sample 'R_hat'
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: v31, points: 20, grade: true, solution: false -->
```python
assert_equal( u_exp(0.5*np.ones(2) ,1,(np.ones(50),np.ones(50)),2),1)
assert_equal( u_exp(0.5*np.ones(2) ,1,(2*np.ones(50),2*np.ones(50)),1.5),2**0.5)

```

- 3.2. Write a function that approximates the solution of the optimal investment problem in the market defined in Question 2, for an investor as described in 3.1..  The function should return a couple, the strategy (as a vector of weights for the two risky assets) and the value of the utility for the optimal portfolio. If no solution is found, 'None' should be returned. [20 points]

<!-- grade_id: q32, points: 0, grade: false, solution: true -->
```python
def optimal_investment(w0,m1,m2,sigma1,sigma2,rho,p):  # modify function name and parameters
    """
    Return pi_hat and expected utility of optimal strategy for investor with utility function with param 'p', and market with params. 'm1,m2,sigma1,sigma2,rho'
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: v32, points: 20, grade: true, solution: false -->
```python

assert_equal(len(optimal_investment(100,0.01,0.01,0.1,0.2,0,1.5)[0]),2)
aux_f = lambda x: np.append(x[0],x[1])
np.testing.assert_allclose(aux_f(optimal_investment(100,0.01,0.02,0.16,0.16,-0.5,1.5)), np.array([2.75, 3.25,10.5 ]), rtol=5e-1, atol=2e-1)




```